cnf(c_0_284,plain-derived,
        ( ~ bird(X2)
        | ~ snail(X1)
        | ~ grain(X1) ),
        inference(pm,[status(thm)],[c_0_277,c_0_12,theory(equality)]),
        [interestingness(0.818210)]).
 
cnf(c_0_532,plain-derived,
        ( ~ wolf(X1)
        | ~ grain(X2)
        | ~ fox(a_bird)
        | ~ plant(X2) ),
        inference(pm,[status(thm)],[c_0_523,c_0_8,theory(equality)]),
        [interestingness(0.814391)]).
 
cnf(c_0_168,conjecture-derived,
        ( eats(X1,X2)
        | ~ snail(X1)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ wolf(a_bird)
        | ~ caterpillar(X2)
        | ~ much_smaller(X2,X1) ),
        inference(pm,[status(thm)],[c_0_159,c_0_8,theory(equality)]),
        [interestingness(0.575651)]).
 
cnf(c_0_161,conjecture-derived,
        ( eats(X2,X3)
        | ~ bird(X1)
        | ~ snail(X2)
        | ~ animal(X3)
        | ~ animal(X2)
        | ~ wolf(X1)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X2) ),
        inference(pm,[status(thm)],[c_0_157,c_0_3,theory(equality)]),
        [interestingness(0.530189)]).
 
cnf(c_0_175,conjecture-derived,
        ( eats(X2,X1)
        | ~ snail(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ wolf(a_bird)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_168,c_0_14,theory(equality)]),
        [interestingness(0.527706)]).
 
cnf(c_0_135,conjecture-derived,
        ( eats(X2,X1)
        | eats(X3,X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ wolf(X3)
        | ~ caterpillar(X1)
        | ~ grain(X4)
        | ~ much_smaller(X2,X3)
        | ~ wolf(X2)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_123,c_0_17,theory(equality)]),
        [interestingness(0.512466)]).
 
cnf(c_0_283,plain-derived,
        ( ~ snail(X1)
        | ~ plant(X1) ),
        inference(pm,[status(thm)],[c_0_277,c_0_8,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_541,plain-derived,
        ( ~ fox(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_537,c_0_11,theory(equality)]),
        [interestingness(0.490847)]).
 
cnf(c_0_147,conjecture-derived,
        ( eats(X2,X1)
        | eats(X3,X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ wolf(X3)
        | ~ caterpillar(X1)
        | ~ much_smaller(X2,X3)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_134,c_0_14,theory(equality)]),
        [interestingness(0.468249)]).
 
cnf(c_0_85,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ caterpillar(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_79,c_0_16,theory(equality)]),
        [interestingness(0.462346)]).
 
cnf(c_0_144,conjecture-derived,
        ( eats(X1,X3)
        | eats(X2,X1)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ wolf(X2)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X1)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_134,c_0_17,theory(equality)]),
        [interestingness(0.448520)]).
 
cnf(c_0_84,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ caterpillar(X1)
        | ~ wolf(X2)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_79,c_0_17,theory(equality)]),
        [interestingness(0.421790)]).
 
