cnf(c_0_281,plain-derived,
        ( ~ wolf(a_bird)
        | ~ grain(X1) ),
        inference(pm,[status(thm)],[c_0_278,c_0_8,theory(equality)]),
        [interestingness(0.666667)]).
 
cnf(c_0_243,plain-derived,
        ( ~ bird(X1)
        | ~ snail(X2)
        | ~ animal(X1)
        | ~ plant(X2) ),
        inference(pm,[status(thm)],[c_0_21,c_0_236,theory(equality)]),
        [interestingness(0.629935)]).
 
cnf(c_0_282,plain-derived,
        ( ~ wolf(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_278,c_0_11,theory(equality)]),
        [interestingness(0.563318)]).
 
cnf(c_0_861,plain-derived,
        ( ~ fox(X1)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_837,c_0_236,theory(equality)]),
        [interestingness(0.546716)]).
 
cnf(c_0_165,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X3)
        | ~ fox(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ caterpillar(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_154,c_0_16,theory(equality)]),
        [interestingness(0.522703)]).
 
cnf(c_0_166,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X3)
        | ~ fox(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_154,c_0_14,theory(equality)]),
        [interestingness(0.481036)]).
 
cnf(c_0_164,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X3)
        | ~ fox(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ caterpillar(X1)
        | ~ wolf(X2)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_154,c_0_17,theory(equality)]),
        [interestingness(0.441264)]).
 
cnf(c_0_163,conjecture-derived,
        ( eats(X1,X2)
        | ~ fox(X1)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ caterpillar(X2)
        | ~ much_smaller(X2,X1) ),
        inference(pm,[status(thm)],[c_0_154,c_0_6,theory(equality)]),
        [interestingness(0.397284)]).
 
cnf(c_0_236,plain-derived,
        ( eats(X1,X2)
        | ~ bird(X1)
        | ~ animal(X1)
        | ~ plant(X2) ),
        inference(pm,[status(thm)],[c_0_231,c_0_10,theory(equality)]),
        [interestingness(0.337340)]).
 
