cnf(c_0_174,conjecture-derived,
        ( eats(X1,X3)
        | eats(X3,X4)
        | ~ wolf(X1)
        | ~ fox(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ animal(X4)
        | ~ caterpillar(X4)
        | ~ much_smaller(X4,X3)
        | ~ fox(X3) ),
        inference(pm,[status(thm)],[c_0_18,c_0_155]),
        [interestingness(0.878115)]).
 
cnf(c_0_383,conjecture-derived,
        ( eats(X2,X3)
        | ~ grain(X1)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X2) ),
        inference(pm,[status(thm)],[c_0_375,c_0_12]),
        [interestingness(0.711638)]).
 
cnf(c_0_393,conjecture-derived,
        ( eats(X2,X1)
        | ~ grain(X3)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_383,c_0_14]),
        [interestingness(0.706757)]).
 
cnf(c_0_101,conjecture-derived,
        ( ~ wolf(X1)
        | ~ grain(X2)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ caterpillar(X2)
        | ~ fox(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_19,c_0_97]),
        [interestingness(0.686787)]).
 
cnf(c_0_96,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ wolf(X2)
        | ~ caterpillar(X1)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_91,c_0_17]),
        [interestingness(0.644758)]).
 
cnf(c_0_173,conjecture-derived,
        ( eats(X1,X3)
        | eats(X2,X4)
        | ~ wolf(X1)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ plant(X3)
        | ~ animal(X4)
        | ~ caterpillar(X4)
        | ~ much_smaller(X4,X2) ),
        inference(pm,[status(thm)],[c_0_18,c_0_155]),
        [interestingness(0.643951)]).
 
cnf(c_0_61,conjecture-derived,
        ( ~ bird(caterpillar_food_of(X1))
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_47,c_0_3]),
        [interestingness(0.634705)]).
 
cnf(c_0_305,conjecture-derived,
        ( eats(X1,X2)
        | ~ fox(X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ caterpillar(X2)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_298,c_0_11]),
        [interestingness(0.629344)]).
 
cnf(c_0_198,conjecture-derived,
        ( eats(X1,X2)
        | ~ fox(X1)
        | ~ animal(X1)
        | ~ plant(X1)
        | ~ animal(X2)
        | ~ caterpillar(X2)
        | ~ much_smaller(X2,X1) ),
        inference(pm,[status(thm)],[c_0_189,c_0_6]),
        [interestingness(0.591936)]).
 
cnf(c_0_64,conjecture-derived,
        ( ~ bird(snail_food_of(X1))
        | ~ snail(X1)
        | ~ wolf(snail_food_of(X1)) ),
        inference(pm,[status(thm)],[c_0_48,c_0_1]),
        [interestingness(0.591667)]).
 
cnf(c_0_65,conjecture-derived,
        ( ~ bird(snail_food_of(X1))
        | ~ snail(X1)
        | ~ fox(snail_food_of(X1)) ),
        inference(pm,[status(thm)],[c_0_48,c_0_2]),
        [interestingness(0.587897)]).
 
cnf(c_0_98,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ wolf(X2)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_91,c_0_14]),
        [interestingness(0.580845)]).
 
cnf(c_0_205,conjecture-derived,
        ( eats(X2,X1)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ plant(X2)
        | ~ animal(X1)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_198,c_0_14]),
        [interestingness(0.539320)]).
 
cnf(c_0_164,conjecture-derived,
        ( eats(X1,X3)
        | eats(X3,X4)
        | ~ wolf(X1)
        | ~ grain(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ animal(X4)
        | ~ caterpillar(X4)
        | ~ much_smaller(X4,X3) ),
        inference(pm,[status(thm)],[c_0_19,c_0_142]),
        [interestingness(0.505241)]).
 
cnf(c_0_69,conjecture-derived,
        ( ~ bird(snail_food_of(a_snail)) ),
        inference(pm,[status(thm)],[c_0_66,c_0_10]),
        [interestingness(0.500000)]).
 
