cnf(c_0_303,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ animal(X3)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X1)
        | ~ grain(X4)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_289,c_0_16]),
        [interestingness(0.928722)]).
 
cnf(c_0_294,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ plant(X4)
        | ~ animal(X3)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_253,c_0_16]),
        [interestingness(0.896709)]).
 
cnf(c_0_254,conjecture-derived,
        ( eats(X3,X1)
        | eats(X1,X4)
        | ~ animal(X1)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ plant(X2)
        | ~ much_smaller(X1,X3)
        | ~ animal(X4)
        | ~ bird(X4)
        | ~ caterpillar(X5)
        | ~ much_smaller(X4,X1) ),
        inference(pm,[status(thm)],[c_0_230,c_0_84]),
        [interestingness(0.865459)]).
 
cnf(c_0_253,conjecture-derived,
        ( eats(X3,X1)
        | eats(X1,X4)
        | ~ animal(X1)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ plant(X2)
        | ~ much_smaller(X1,X3)
        | ~ animal(X4)
        | ~ caterpillar(X4)
        | ~ much_smaller(X4,X1) ),
        inference(pm,[status(thm)],[c_0_230,c_0_41]),
        [interestingness(0.840221)]).
 
cnf(c_0_306,conjecture-derived,
        ( eats(X3,X2)
        | eats(X2,X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ caterpillar(X1)
        | ~ grain(X4)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_289,c_0_14]),
        [interestingness(0.803722)]).
 
cnf(c_0_292,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ plant(X4)
        | ~ animal(X3)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X1)
        | ~ wolf(X2)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_253,c_0_17]),
        [interestingness(0.789661)]).
 
cnf(c_0_293,conjecture-derived,
        ( eats(X3,X2)
        | eats(X2,X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ plant(X4)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ caterpillar(X1)
        | ~ wolf(X2)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_253,c_0_17]),
        [interestingness(0.789661)]).
 
cnf(c_0_430,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ animal(X3)
        | ~ plant(snail_food_of(X3))
        | ~ much_smaller(X3,X1)
        | ~ snail(X3)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_409,c_0_14]),
        [interestingness(0.777663)]).
 
cnf(c_0_295,conjecture-derived,
        ( eats(X3,X2)
        | eats(X2,X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ plant(X4)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ caterpillar(X1)
        | ~ fox(X2)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_253,c_0_16]),
        [interestingness(0.771709)]).
 
cnf(c_0_296,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ plant(X4)
        | ~ animal(X3)
        | ~ caterpillar(X3)
        | ~ much_smaller(X3,X1)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_253,c_0_14]),
        [interestingness(0.771709)]).
 
cnf(c_0_244,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ plant(caterpillar_food_of(X1))
        | ~ much_smaller(X1,X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_230,c_0_23]),
        [interestingness(0.655465)]).
 
cnf(c_0_251,conjecture-derived,
        ( eats(X3,X1)
        | ~ animal(X1)
        | ~ animal(X3)
        | ~ fox(X3)
        | ~ plant(X2)
        | ~ much_smaller(X1,X3)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_230,c_0_77]),
        [interestingness(0.654522)]).
 
cnf(c_0_245,conjecture-derived,
        ( eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ fox(X2)
        | ~ plant(snail_food_of(X1))
        | ~ much_smaller(X1,X2)
        | ~ snail(X1) ),
        inference(pm,[status(thm)],[c_0_230,c_0_25]),
        [interestingness(0.581731)]).
 
