cnf(c_0_281,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ wolf(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ bird(X3)
        | ~ much_smaller(X3,X1)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_272,c_0_17,theory(equality)]),
        [interestingness(0.785068)]).
 
cnf(c_0_544,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ animal(X3)
        | ~ bird(X3)
        | ~ much_smaller(X3,X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_533,c_0_16,theory(equality)]),
        [interestingness(0.776520)]).
 
cnf(c_0_242,conjecture-derived,
        ( eats(X2,X1)
        | eats(X2,X3)
        | eats(X1,X4)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ plant(X3)
        | ~ animal(X4)
        | ~ bird(X4)
        | ~ much_smaller(X4,X1)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_225,c_0_14,theory(equality)]),
        [interestingness(0.771609)]).
 
cnf(c_0_243,conjecture-derived,
        ( eats(X3,X2)
        | eats(X3,X4)
        | eats(X2,X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ plant(X4)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ bird(X1)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_225,c_0_14,theory(equality)]),
        [interestingness(0.771609)]).
 
cnf(c_0_159,conjecture-derived,
        ( eats(X1,X3)
        | ~ bird(X1)
        | ~ snail(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ bird(X3)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1) ),
        inference(pm,[status(thm)],[c_0_21,c_0_145,theory(equality)]),
        [interestingness(0.728102)]).
 
cnf(c_0_177,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X2)
        | ~ grain(X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_167,c_0_14,theory(equality)]),
        [interestingness(0.709762)]).
 
cnf(c_0_165,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X2)
        | ~ fox(X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ plant(X3)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_157,c_0_14,theory(equality)]),
        [interestingness(0.698533)]).
 
cnf(c_0_284,conjecture-derived,
        ( eats(X3,X2)
        | eats(X2,X1)
        | ~ wolf(X3)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ bird(X1)
        | ~ fox(X2) ),
        inference(pm,[status(thm)],[c_0_272,c_0_16,theory(equality)]),
        [interestingness(0.691318)]).
 
cnf(c_0_154,conjecture-derived,
        ( eats(X2,X3)
        | eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ plant(X3)
        | ~ fox(X2) ),
        inference(pm,[status(thm)],[c_0_145,c_0_16,theory(equality)]),
        [interestingness(0.690084)]).
 
cnf(c_0_545,conjecture-derived,
        ( eats(X3,X2)
        | eats(X2,X1)
        | ~ fox(X3)
        | ~ animal(X3)
        | ~ animal(X2)
        | ~ much_smaller(X2,X3)
        | ~ animal(X1)
        | ~ bird(X1)
        | ~ fox(X2) ),
        inference(pm,[status(thm)],[c_0_533,c_0_16,theory(equality)]),
        [interestingness(0.682770)]).
 
cnf(c_0_285,conjecture-derived,
        ( eats(X2,X1)
        | eats(X1,X3)
        | ~ wolf(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ animal(X3)
        | ~ bird(X3)
        | ~ much_smaller(X3,X1)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_272,c_0_14,theory(equality)]),
        [interestingness(0.675150)]).
 
cnf(c_0_155,conjecture-derived,
        ( eats(X2,X3)
        | eats(X2,X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ plant(X3)
        | ~ bird(X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_145,c_0_14,theory(equality)]),
        [interestingness(0.668216)]).
 
cnf(c_0_179,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_174,c_0_17,theory(equality)]),
        [interestingness(0.640677)]).
 
cnf(c_0_174,conjecture-derived,
        ( eats(X1,X2)
        | ~ wolf(X1)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ bird(X2)
        | ~ much_smaller(X2,X1) ),
        inference(pm,[status(thm)],[c_0_167,c_0_11,theory(equality)]),
        [interestingness(0.583530)]).
 
cnf(c_0_180,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X2)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ fox(X2) ),
        inference(pm,[status(thm)],[c_0_174,c_0_16,theory(equality)]),
        [interestingness(0.578177)]).
 
cnf(c_0_247,conjecture-derived,
        ( eats(X1,X3)
        | eats(X3,X4)
        | ~ wolf(X1)
        | ~ grain(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ animal(X4)
        | ~ bird(X4)
        | ~ much_smaller(X4,X3) ),
        inference(pm,[status(thm)],[c_0_19,c_0_225,theory(equality)]),
        [interestingness(0.547919)]).
 
