cnf(c_0_74,conjecture-derived,
        ( eats(X1,X3)
        | ~ wolf(X1)
        | ~ grain(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ caterpillar(X3) ),
        inference(pm,[status(thm)],[c_0_19,c_0_60,theory(equality)]),
        [interestingness(0.632679)]).
 
cnf(c_0_71,conjecture-derived,
        ( eats(X2,X1)
        | eats(X2,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ plant(X3)
        | ~ caterpillar(X1)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_60,c_0_14,theory(equality)]),
        [interestingness(0.631995)]).
 
cnf(c_0_75,conjecture-derived,
        ( eats(X1,X3)
        | ~ bird(X1)
        | ~ snail(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ caterpillar(X3) ),
        inference(pm,[status(thm)],[c_0_21,c_0_60,theory(equality)]),
        [interestingness(0.627381)]).
 
cnf(c_0_73,conjecture-derived,
        ( eats(X1,X3)
        | ~ wolf(X1)
        | ~ fox(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ caterpillar(X3) ),
        inference(pm,[status(thm)],[c_0_18,c_0_60,theory(equality)]),
        [interestingness(0.626334)]).
 
cnf(c_0_35,conjecture-derived,
        ( ~ snail(X1)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_34,c_0_8,theory(equality)]),
        [interestingness(0.625758)]).
 
cnf(c_0_1675,conjecture-derived,
        ( ~ bird(X1)
        | ~ grain(X1) ),
        inference(pm,[status(thm)],[c_0_1670,c_0_3,theory(equality)]),
        [interestingness(0.625758)]).
 
cnf(c_0_379,conjecture-derived,
        ( ~ snail(X1)
        | ~ grain(X1) ),
        inference(pm,[status(thm)],[c_0_375,c_0_12,theory(equality)]),
        [interestingness(0.625000)]).
 
cnf(c_0_218,conjecture-derived,
        ( ~ fox(X1)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_214,c_0_11,theory(equality)]),
        [interestingness(0.623737)]).
 
cnf(c_0_672,conjecture-derived,
        ( ~ wolf(a_fox)
        | ~ fox(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_671,c_0_7,theory(equality)]),
        [interestingness(0.580773)]).
 
cnf(c_0_1666,conjecture-derived,
        ( ~ plant(X1)
        | ~ animal(X1)
        | ~ bird(X1) ),
        inference(pm,[status(thm)],[c_0_1659,c_0_7,theory(equality)]),
        [interestingness(0.549612)]).
 
cnf(c_0_37,conjecture-derived,
        ( ~ caterpillar(a_snail) ),
        inference(pm,[status(thm)],[c_0_35,c_0_10,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_173,conjecture-derived,
        ( ~ fox(a_caterpillar) ),
        inference(pm,[status(thm)],[c_0_170,c_0_11,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_332,conjecture-derived,
        ( ~ snail(a_grain) ),
        inference(pm,[status(thm)],[c_0_331,c_0_11,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_811,conjecture-derived,
        ( ~ fox(a_bird) ),
        inference(pm,[status(thm)],[c_0_808,c_0_11,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_1678,conjecture-derived,
        ( ~ bird(a_grain) ),
        inference(pm,[status(thm)],[c_0_1675,c_0_11,theory(equality)]),
        [interestingness(0.500000)]).
 
cnf(c_0_60,conjecture-derived,
        ( eats(X2,X1)
        | eats(X2,X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ plant(X3)
        | ~ much_smaller(X1,X2)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_47,c_0_22,theory(equality)]),
        [interestingness(0.427258)]).
 
