cnf(c_0_154,conjecture-derived,
        ( eats(X1,X3)
        | ~ bird(X1)
        | ~ snail(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X4)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ bird(X3) ),
        inference(pm,[status(thm)],[c_0_21,c_0_140]),
        [interestingness(0.975848)]).
 
cnf(c_0_147,conjecture-derived,
        ( eats(X2,X3)
        | eats(X2,X4)
        | ~ animal(X3)
        | ~ animal(X2)
        | ~ plant(X4)
        | ~ much_smaller(X3,X2)
        | ~ bird(X3)
        | ~ snail(X1) ),
        inference(pm,[status(thm)],[c_0_140,c_0_24]),
        [interestingness(0.975848)]).
 
cnf(c_0_172,conjecture-derived,
        ( eats(X1,X3)
        | ~ wolf(X1)
        | ~ fox(X2)
        | ~ animal(X2)
        | ~ animal(X1)
        | ~ plant(X3)
        | ~ bird(X2)
        | ~ fox(X1)
        | ~ grain(X4) ),
        inference(pm,[status(thm)],[c_0_18,c_0_158]),
        [interestingness(0.975144)]).
 
cnf(c_0_153,conjecture-derived,
        ( eats(X1,X3)
        | ~ wolf(X1)
        | ~ grain(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X4)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ bird(X3) ),
        inference(pm,[status(thm)],[c_0_19,c_0_140]),
        [interestingness(0.975065)]).
 
cnf(c_0_245,conjecture-derived,
        ( eats(X2,X1)
        | ~ wolf(X2)
        | ~ grain(X3)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X1)
        | ~ fox(X1) ),
        inference(pm,[status(thm)],[c_0_237,c_0_17]),
        [interestingness(0.959098)]).
 
cnf(c_0_152,conjecture-derived,
        ( eats(X1,X3)
        | ~ wolf(X1)
        | ~ fox(X2)
        | ~ animal(X3)
        | ~ animal(X1)
        | ~ plant(X4)
        | ~ plant(X2)
        | ~ much_smaller(X3,X1)
        | ~ bird(X3) ),
        inference(pm,[status(thm)],[c_0_18,c_0_140]),
        [interestingness(0.946594)]).
 
cnf(c_0_146,conjecture-derived,
        ( eats(X2,X3)
        | eats(X2,X4)
        | ~ animal(X3)
        | ~ animal(X2)
        | ~ plant(X4)
        | ~ much_smaller(X3,X2)
        | ~ bird(X3)
        | ~ caterpillar(X1) ),
        inference(pm,[status(thm)],[c_0_140,c_0_22]),
        [interestingness(0.911942)]).
 
cnf(c_0_494,conjecture-derived,
        ( ~ wolf(X1)
        | ~ grain(X2)
        | ~ fox(X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_19,c_0_488]),
        [interestingness(0.820729)]).
 
cnf(c_0_145,conjecture-derived,
        ( eats(X2,X3)
        | eats(X2,X4)
        | ~ animal(X3)
        | ~ animal(X2)
        | ~ plant(X4)
        | ~ much_smaller(X3,X2)
        | ~ bird(X3)
        | ~ grain(X1) ),
        inference(pm,[status(thm)],[c_0_140,c_0_12]),
        [interestingness(0.694515)]).
 
cnf(c_0_488,conjecture-derived,
        ( eats(X1,X2)
        | ~ fox(X1)
        | ~ animal(X1)
        | ~ animal(X2)
        | ~ bird(X2) ),
        inference(pm,[status(thm)],[c_0_481,c_0_11]),
        [interestingness(0.691331)]).
 
